---
title: 'Chapter 1'
description: 'Chapter description goes here.'
free_preview: true
---

## What is a program?

```yaml
type: VideoExercise
key: 3b396fe2fb
xp: 50
```

`@projector_key`
b04e1e9d05eb9540d3212141da9d42c1

---

## Using the IPython shell

```yaml
type: MultipleChoiceExercise
key: d81817c08a
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Using Python as a calculator

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
