---
title: 'What is a program?'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## What is a program?

```yaml
type: VideoExercise
key: 3b396fe2fb
xp: 50
```

`@projector_key`
b04e1e9d05eb9540d3212141da9d42c1
